---
rule_files:
  - "zones_changes.rule"
evaluation_interval: 1m
tests:
  # HighAmountCollateralsUnderRisk
  - interval: 5m
    input_series:
      - series: maker_risks_collateral_percentage{zone="D"}
        values: 20+0x10
    alert_rule_test:
      # no fire
      - eval_time: 5m
        alertname: HighAmountCollateralsUnderRisk

      - eval_time: 30m
        alertname: HighAmountCollateralsUnderRisk
        exp_alerts:
          - exp_labels:
              severity: high
            exp_annotations:
              summary: High percent of Maker collaterals in dangerous zone 🔥
              description: It is 20.0% of collaterals on Maker in risky zone!

  # MakerCollateralsLiquidation
  - interval: 5m
    input_series:
      - series: maker_risks_collateral_percentage{zone="liquidation"}
        values: 0.01+0x10
    alert_rule_test:
      - eval_time: 5m
        alertname: MakerCollateralsLiquidation
        exp_alerts:
          - exp_labels:
              severity: high
              zone: liquidation
            exp_annotations:
              summary: Maker collaterals under liquidation has been detected!

# vim: set ts=2 sw=2 ft=yaml:
